Avian-specific Salmonella transition to endemicity is accompanied by localized resistome and mobilome interaction

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary:

The investigators in this study analyzed the dataset assembly from 540 Salmonella isolates, and those from 45 recent isolates from Zhejiang University of China. The analysis and comparison of the resistome and mobilome of these isolates identified a significantly higher rate of cross-region dissemination compared to localized propagation. This study highlights the key role of the resistome in driving the transition and evolutionary history of S. Gallinarum.

Strengths:

The isolates included in this study were from 16 countries in the past century (1920 to 2023). While the study uses S. Gallinarun as the prototype, the conclusion from this work will likely apply to other Salmonella serotypes and other pathogens.

Thank you very much for your positive feedback. We recognize, as you noted, that emphasizing Salmonella enterica Serovar Gallinarum in the title may lead readers to perceive our methods and conclusions as overly restrictive. In light of your evaluation of our work, we have revised the title to: “Avian-specific Salmonella transition to endemicity is accompanied by localized resistome and mobilome interaction” We believe this final version not only reflects the applicability of our conclusions, as you appreciated, but also addresses your previous suggestion to highlight the resistome and mobilome.

Revisions in the manuscript Lines: 1-3

Weaknesses:

While the isolates came from 16 countries, most strains in this study were originally from China.

We believe that this issue was discussed in detail in our previous response. Although potential bias exists, we have minimized its impact by constructing the largest global S. Gallinarum genome dataset to date. In addition, we have further emphasized these limitations in the manuscript.

Comments on revisions:

This reviewer is happy with the detailed responses from the authors regarding revising this manuscript. I do not have further comments.

We greatly appreciate your positive feedback and are pleased that our responses have addressed your concerns.

Reviewer #2 (Public review):

Summary:

The authors sequence 45 new samples of S. Gallinarum, a commensal Salmonella found in chickens, which can sometimes cause disease. They combine these sequences with around 500 from public databases, determine the population structure of the pathogen, and coarse relationships of lineages with geography. The authors further investigate known anti-microbial genes found in these genomes, how they associate with each other, whether they have been horizontally transferred, and date the emergence of clades.

Strengths:

- It doesn't seem that much is known about this serovar, so publicly available new sequences from a high burden region are a valuable addition to the literature.

- Combining these sequences with publicly available sequences is a good way to better contextualise any findings.

- The genomic analyses have been greatly improved since the first version of the manuscript, and appropriately analyse the population and date emergence of clades.

- The SNP thresholds are contextualised in terms of evolutionary time.

- The importance and context of the findings are fairly well described.

Thank you so much for your thorough review and constructive comments on the manuscript.

Weaknesses:

- There are still a few issues with the genomic analyses, although they no longer undermine the main conclusions:

We are grateful for the valuable time and effort you have dedicated to improving our manuscript. In this revision, we have provided a point-by-point response to each of your concerns. Moreover, with the addition of new supplementary materials and modifications to the figures, we have re-examined and adjusted the numbering of figures and supplementary materials in the text to ensure they appear correctly in the manuscript.

(1) Although the SNP distance is now considered in terms of time, the 5 SNP distance presented still represents ~7yrs evolution, so it is unlikely to be a transmission event, as described. It would be better to use a much lower threshold or describe the interpretation of these clusters more clearly. Bringing in epidemiological evidence or external references on the likely time interval between transmissions would be helpful.

We sincerely thank you for highlighting this issue. We appreciate your concern regarding the use of a 5-SNP threshold to define a transmission event, especially given the approximate 7-year evolutionary timeframe. Considering our updated estimate for the evolutionary rate of S. Gallinarum (approximately 0.74 SNPs per year, with a 95% HPD range of 0.42 to 1.06), we have revised the manuscript to use a 2-SNP threshold (approximately representing less than two years of evolution) to better control the temporal span of transmission events. In addition, we have updated the manuscript to reflect this new threshold and demonstrated that the use of a more stringent SNP threshold does not affect the overall conclusions of the study.

Specifically, we adopted the newly established 2-SNP threshold to update Figure 3a and corresponding Supplementary Figure 8. The heatmap on the far right of New Figure 3a illustrates the SNP distances among 45 newly isolated S. Gallinarum strains from two locations in Zhejiang Province (Taishun and Yueqing). New Supplementary Figure 8 simulates potential transmission events between the bvSP strains isolated from Zhejiang Province (n=95) and those from other regions of China with available provincial information (n=435). These analyses collectively demonstrate the localized transmission patterns of bvSP within China.

For New Figure 3a, we found that even with the 2-SNP threshold, the number of potential transmission events among the 45 newly isolated S. Gallinarum strains from the two Zhejiang locations (Taishun and Yueqing) remains unchanged. In fact, we observed that the results from SNP tracing using an SNP threshold of less than 5 are consistent (see Author response image 1).

Clustering results of 45 newly isolated S.

Gallinarum strains using different SNP thresholds of 1, 2, 3, 4, and 5 SNPs. The five subplots represent the clustering results under each threshold. Each point corresponds to an individual strain, and lines connect strains with potential transmission relationships.

For New Supplementary Figure 8, we employed the 2-SNP threshold and found that the number of transmission events between the bvSP strains isolated from Zhejiang Province (n=95) and those from other Chinese provinces (n=435) decreased from 91 to 53. The names of the strains involved in these potential transmission events are listed in Supplementary Table 5.

Revisions in the manuscript

Lines: 352-357

Figures: Figure 3; Supplementary Figure 8

Table: Supplementary Table 5

(2) The HGT definition has not fundamentally been changed and therefore still has some issues, mainly that vertical evolution is still not systematically controlled for.

We sincerely thank you for highlighting this issue. We hope the following explanation will help clarify and improve our manuscript, as well as address your concerns.

In bacteria, mobile genetic elements (MGEs) such as plasmids, transposons, integrons, and prophages, as mentioned in our manuscript, are segments of DNA that encode enzymes and proteins responsible for mediating the movement of genetic material between bacterial genomes (commonly referred to as “jumping genes”). These MGEs contribute to the mechanisms of horizontal gene transfer (HGT) in Salmonella, including transduction (via prophages), conjugation (via plasmids), and transposition (via integrons and transposons) (Nat Rev Microbiol. 2005 Sep;3(9):722-32). These “jumping genes” can enable Salmonella to acquire additional antimicrobial resistance genes (ARGs), which may not only originate from other Salmonella strains but also from distantly related species.

To further address your concern regarding the systematic control of vertical evolution, we employed the HGTphyloDetect pipeline developed by Le Yuan et al. (Brief Bioinform. 2023 Mar 19;24(2):bbad035) to control for vertical evolution in the ARG sequences mentioned in our manuscript. We chose HGTphyloDetect because, as noted, "jumping genes" often occur among evolutionarily distant species, rendering the use of Gubbins potentially unsuitable for these distant HGT events.

Using the HGTphyloDetect pipeline, we extracted base sequences for the eight ARGs shown in Figure 6b with an HGT frequency greater than zero (blaTEM-1B, sul1, dfrA17, aadA5, sul2, aph(3’’)-Ib, tet(A), aph(6)-Id). For blaTEM-1B, sul1, dfrA17, aadA5, and sul2, the HGT frequency reached 100% across different isolates, indicating that these ARG sequences have a unique sequence type. In contrast, due to the ResFinder settings requiring both similarity and coverage to meet a minimum value of 90%, the base sequences for aph(3’’)-Ib, tet(A), and aph(6)-Id are not unique. Consequently, we applied the HGTphyloDetect pipeline individually to each sequence type of ARGs to verify their association with HGT events. Specifically, among 436 bvSP isolates collected in China, we identified two sequence types of aph(3’’)-Ib, four sequence types of tet(A), and three sequence types of aph(6)-Id.

Subsequently, to identify potential ARGs horizontally acquired from evolutionarily distant organisms, we queried the translated amino acid sequences of each ARG against the National Center for Biotechnology Information (NCBI) non-redundant protein database. We then evaluated whether these sequences were products of HGT by calculating Alien Index (AI) scores and out_perc values.

The calculation of AI score is as follows:

In this study, bbhG and bbhO represent the E-values of the best blast hit in ingroup and outgroup lineages, respectively. The outgroup lineage is defined as all species outside of the kingdom, while the ingroup lineage encompasses species within the kingdom but outside of the subphylum. An AI score ≥ 45 is considered a strong indicator that the gene in question is likely derived from an HGT event.

Regarding the calculation method for out_perc:

Finally, according to the definition provided by the HGTphyloDetect pipeline, ARGs with AI score ≥ 45 and out_perc ≥ 90% are presumed to be potential candidates for HGT from evolutionarily distant species. We have compiled the calculation results for the aforementioned genes in New Supplementary Table 9. The results indicate that all ARGs presented in Figure 6b, which exhibited a HGT frequency greater than zero, were acquired horizontally by S. Gallinarum. Based on these findings, we have revised the manuscript accordingly.

Revisions in the manuscript

Lines: 302-307; 616-650; 955-957

Table: Supplementary Table 9

Using a 5kb window is not sufficient, as LD may extend across the entire genome.

We agree with your point that linkage disequilibrium (LD) could influence the transmission of genes within chromosomal regions. LD can lead to the non-random cooccurrence of alleles at different loci within a population. Considering that horizontal gene transfer (HGT) events involving more distantly related ARGs may be accompanied by vertical propagation on chromosomes, and to simultaneously assess the impact of LD, we conducted two evaluations.

It is important to note that the following assessments are based on the assumption that plasmid replicons detected by PlasmidsFinder are part of self-replicating, extrachromosomal DNA.

(1) In the revised pipeline used to calculate ARG HGT frequencies, we categorized a total of 621 ARGs carried by 436 bvSP isolates collected in China and found that 415 of these ARGs were located on MGEs. We further investigated the distribution of these 415 ARGs across different MGEs, taking into account the complex nesting relationships among them. We observed that 90% of the ARGs (372/415) were located on plasmid contigs. It is important to clarify that this finding does not contradict our statement in the manuscript regarding plasmids and transposons as the primary reservoirs for resistome geo-temporal dissemination. This is because transposons, integrons, and prophages carrying ARGs can also be found on plasmids. Additionally, only 25 bvSG isolates from China contained ARGs, which were likely acquired via transposons or integrons located on the chromosome.

(2) In our manuscript, we searched for ARGs within a 5kb upstream and downstream region (a total of 10kb) of transposons and integrons (The BLASTn parameters used in the Bacant pipeline to identify transposons and integrons were set to a coverage threshold of 60%, rather than 100%). However, in light of the potential impact of LD on vertical transmission, we expanded our search to include a 10kb upstream and downstream range (a total of 20kb) for these 25 isolates. The decision to expand the search range to 10kb upstream and downstream range is based on the following two considerations: 1) Based on literature, we determined the overall lengths of the integrons and transposons carried by the 25 isolates (Tn801, Tn6205, Tn1721, In498, In1440, In473, and In282), and found that the maximum length of these elements is ~13.5 kb. Using a 10kb upstream and downstream threshold effectively covers these integrons/transposons. 2) The limitation posed by genomic fragmentation due to next-generation sequencing, which restrict the search range. We present the results of this expanded search for colocalization of ARGs with transposons and integrons at: Figshare: https://doi.org/10.6084/m9.figshare.28129130.v1

We found that these results were consistent with those obtained using the previous search range.

Taken together, these results suggest that although linkage disequilibrium may influence genetic processes within chromosomal regions—particularly for the few chromosomeassociated antibiotic resistance genes linked to integrons and transposons—the overall impact in our study is likely minimal. This conclusion is supported by the observation that 90% of the ARGs in our dataset are located on plasmids, and even an expanded search range does not alter this outcome. Additionally, by incorporating Alien Index scores and calculating out_perc, we can further confirm the occurrence of horizontal gene transfer events.

However, it is undeniable that other studies using our current pipeline may be affected. As a temporary remedial measure, we have included a note in the "README" file as below (https://github.com/tjiaa/Cal_HGT_Frequency):

“Note: Considering that ARGs located on the chromosome and carried by mobile genetic elements—such as integrons and transposons—may introduce potential computational errors, we recommend evaluating the number of ARGs associated with these elements on the chromosome during your analysis. If a majority of ARGs in your dataset fall into this category, we suggest using additional methods to evaluate the potential impact of linkage disequilibrium. Additionally, by modifying the “MGE_start” and “MGE_end” parameters in the “eLife_MGE_ARG_Co_location.ipynb” script, you can assess the distance between different ARGs and integrons or transposons on the chromosome. This approach will further aid in evaluating the impact of linkage disequilibrium on the genetic process.”

We believe this approach will assist researchers in further assessing the potential impact of vertical evolution and help other users determine whether additional methods are necessary to account for such effects.

As the authors have now run gubbins correctly, they could use the results from this existing analysis to find recent HGT.

We sincerely thank you for your valuable suggestion. Utilizing additional methods to predict potential horizontal gene transfer (HGT) events could indeed enhance the robustness of the results. However, "jumping genes" often occur among evolutionarily distant species, rendering the use of Gubbins potentially unsuitable for these distant HGT events.

Furthermore, the primary focus of our study is to identify HGT of antimicrobial resistance genes (ARGs) in the Salmonella genome driven by mobile genetic elements. Therefore, we employed the HGTphyloDetect pipeline developed by Le Yuan et al. (Brief Bioinform. 2023 Mar 19;24(2):bbad035) to control for vertical evolution in the ARG sequences. The specific computational methods and conclusions have been detailed above.

To definite mobilisation, perhaps a standard pipeline such (e.g. https://github.com/EBIMetagenomics/mobilome-annotation-pipeline) would be more convincing.

Thank you for your valuable suggestion. We agree that defining mobilization using a standardized pipeline can add rigor and clarity to our analysis. The pipeline you referenced (https://github.com/EBI-Metagenomics/mobilome-annotation-pipeline) is an excellent resource and provides a robust approach to the identification and annotation of mobile genetic elements.

We have examined and run this pipeline, which uses “IntegronFinder” and “ICEfinder” to detect integrons, “geNomad” to identify plasmids, and “geNomad” and “VIRify” to detect prophages. Our initial checks revealed that the numbers of integrons, plasmids, and prophages identified using this pipeline were consistent with those detected in our study. However, due to the significantly different output formats, the results from this pipeline could not be integrated with the pipeline we used for calculating HGT frequency.

We will incorporate the standardized pipeline you suggested in future studies to further improve the reliability of our findings.

(3) The invasiveness index is better described, but the authors still did not provide convincing evidence that the small difference is actually biologically meaningful (there was no statistical difference between the two strains provided in response Figure 6). What do other Salmonella papers using this approach find, and can their links be brought in? If there is still no good evidence, a better description of this difference would help make the conclusions better supported.

We sincerely appreciate your thoughtful feedback. The initial introduction of the invasiveness index in our manuscript aimed to quantitatively assess the differences in invasiveness between two geographically distinct strains of S. Gallinarum (isolated from Taishun and Yueqing) by comparing the degradation of 196 top predicted genes associated with invasiveness in their genomes. We found a highly significant statistical difference (P < 0.0001) in the invasiveness index between them.

Several studies have also employed the invasiveness index to predict biological relevance in Salmonella strains, and we believe these examples provide further context for our approach:

(1) Caisey V. Pulford et al, Nat Microbiol, 2021, used the same method to calculate the invasiveness index for Salmonella Typhimurium and employed it to characterize the invasiveness of different lineage strains. They found that Salmonella in Lineage-3 exhibited the highest invasiveness index, suggesting an adaptation from an intestinal to a systemic lifestyle. The authors noted, "Although the invasiveness index cannot yet be experimentally validated, Salmonella isolates with different invasiveness indices produce distinct clinical symptoms in a human population (BMC Med. 2020 Jul 17; 18(1):212)". They emphasized the necessity of developing more robust methods to measure Salmonella invasiveness.

(2) Sandra Van Puyvelde et al, Nat Commun, 2019, reported that Salmonella Typhimurium sequence type 313 (ST313) lineage II.1 exhibited a higher invasiveness index compared to lineage II, suggesting that the two lineages might have distinct adaptations to an invasive lifestyle. Further experiments demonstrated significant differences between these lineages in terms of biofilm formation (A red dry and rough (RDAR) assay) and metabolic capacity for carbon compounds.

(3) Wim L. Cuypers et al, Nat Commun, 2023, calculated the invasiveness index for 284 global Salmonella Concord strains across different lineages and found that Lineage-4 potentially exhibited the highest invasiveness.

Given these evidences, we acknowledge that no significant difference in mortality was observed between the L2b and L3b S. Gallinarum strains in 16-day-old SPF chicken embryos. Existing literature suggests that strains with higher invasiveness indices may still exhibit differences in biofilm formation and metabolic capacities, reflecting their adaptation to different host environments. As such, we maintain that the invasiveness index remains a valuable metric for evaluating the genomic differences between S. Gallinarum strains from Taishun and Yueqing. We plan to further investigate these differences through phenotypic experiments in our next research.

In the revised manuscript, we have added the following discussion along with additional references:

Lines 358-365: “Moreover, the invasiveness index of bvSP from Taishun and Yueqing suggests that different lineages of S. Gallinarum recovered from distinct regions may exhibit biological differences. Previous studies have shown that strains with higher invasiveness indexes tend to be more virulent in hosts (30, 31), potentially causing neurological or arthritic symptoms in S. Gallinarum infections. Furthermore, strains with varying invasiveness indexes have been confirmed to differ in their biofilm formation abilities and metabolic capacities for carbon compounds (32).”

Revisions in the manuscript:

Lines: 358-365, 806-827.

In summary, the analysis is broadly well described and feels appropriate. Some of the conclusions are still not fully supported, although the main points and context of the paper now appear sound.

Thank you so much for your positive evaluation of our work. We hope that the revised manuscript meets your expectations and offers a more accurate interpretation of our findings.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

This is a great improvement over the first version and I thank the authors for a thorough response, as well as changing their conclusions in response to their improvements.

Other small remaining issues:

Figure 3: Heatmap of SNPs is hard to read in grayscale. It also just represents the between clade distances already shown by the tree. It would be more useful to present intraclade distances only to see the SNP resolution _within_ each lineage. Using a better colour scheme would also help.

Thank you for your insightful comments and suggestions regarding Figure 3. We agree that the grayscale heatmap may present challenges in terms of visual clarity. To address this, we have updated the heatmap with a more distinct color gradient, ensuring better contrast and easier interpretation (New Figure 3).

Regarding your second suggestion: "It would be more useful to present intraclade distances only to see the SNP resolution within each lineage," we believe it is already addressed in the current version of New Figure 3. Specifically, the heatmap on the right side of New Figure 3 illustrates the SNP distances between S. Gallinarum isolates from Taishun and Yueqing, with the goal of demonstrating that genomic variation within isolates from a single region is generally smaller compared to those from different regions. In this figure, 45 newly isolated S. Gallinarum strains are categorized into two lineages: L2b and L3b. The heatmap on the right side of Figure 3 displays the SNP distances between all pairwise combinations of these 45 strains, where the intraclade distances are represented by the red regions (highlighting the pairwise distances within each lineage, specifically L3b and L2b, which are indicated by two triangles). The between-clade distances are shown by the blue regions.

We also believe in further exploring the intraclade distances across the entire dataset of 580 S. Gallinarum strains, as it could provide additional insights. However, this analysis would extend beyond the scope of the current section.

Revisions in the manuscript Line: 998

Figure: Figure 3

Please remove Figure 6c, it does not add anything to the paper and raises questions about performing this regression.

Thank you for pointing out this issue. We have removed Figure 6c and the corresponding description in the "Results" section from the manuscript (New Figure 6).

Revisions in the manuscript Lines: 316, 319, 1035-1041.

Figure: Figure 6

Again, thank you all for your time and efforts in reviewing our work. We believe the improved manuscript meets the high standards of the journal.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Bacterial regional demonstration after global dissemination is an essential pathway for selecting distinct finesses. However, the evolution of the resistome during the transition to endemicity remains unaddressed. Using the most comprehensive whole-genome sequencing dataset of Salmonella enterica serovar Gallinarum (S. Gallinarum) collected from 15 countries, including 45 newly recovered samples from two related local regions, we established the relationship among avian-specific pathogen genetic profiles and localization patterns. Initially, we revealed the international transmission and evolutionary history of S. Gallinarum to recent endemicity through phylogenetic analysis conducted using a spatiotemporal Bayesian framework. Our findings indicate that the independent acquisition of the resistome via the mobilome, primarily through plasmids, transposons, and prophages, shapes a unique antimicrobial resistance profile among different lineages. Notably, the mobilome-resistome combination among distinct lineages exhibits a geographical-specific manner, further supporting a localized endemic mobilome-driven process. Collectively, this study elucidates resistome adaptation in the endemic transition of an avian-specific pathogen, likely driven by the localized farming style, and provides valuable insights for targeted interventions.

Revised version: Bacterial regional demonstration after global dissemination is an essential pathway for selecting distinct finesses. However, the evolution of the resistome during the transition to endemicity remains unaddressed. Using the most comprehensive whole-genome sequencing dataset of Salmonella enterica serovar Gallinarum (S. Gallinarum) collected from 15 countries, including 45 newly recovered samples from two related local regions, we established the relationship among avian-specific pathogen genetic profiles and localization patterns. Initially, we revealed the international transmission and evolutionary history of S. Gallinarum to recent endemicity through phylogenetic analysis conducted using a spatiotemporal Bayesian framework. Our findings indicate that the independent acquisition of the resistome via the mobilome, primarily through plasmids and transposons shapes a unique antimicrobial resistance profile among different lineages. Notably, the mobilome-resistome combination among distinct lineages exhibits a geographical-specific manner, further supporting a localized endemic mobilome-driven process. Collectively, this study elucidates resistome adaptation in the endemic transition of an avian-specific pathogen, likely driven by the localized farming style, and provides valuable insights for targeted interventions.

Introduction

The acquisition of antimicrobial resistance (AMR) by pathogens is well-established as one of the most severe threats of the 21st century (1–3). It is estimated that over 10 million people could die per year due to AMR by 2050 (4). Antimicrobial resistance genes (ARGs), collectively known as the resistome, play a pivotal role in AMR development, progression, and amplification (5). Pathogens typically acquire resistome through mobilomes, like plasmids, via horizontal gene transfer (HGT) (6, 7). From an evolutionary perspective, the dynamic change in the resistome is crucial for enhancing pathogen fitness and enabling it to colonize new ecological niches (8). Key factors contributing to resistome alterations include disease burden, human activities, climate change, and geographical selection forces (9–13). Understanding the evolutionary trajectory of the AMR and related reservoirs is necessary, as disease management strategies will differ substantially (14). However, knowledge regarding the overall profile of pathogen regional adaptation related to the stepwise dynamics of the resistome remains limited.

Earlier version: Before recent advancements in whole-genome sequencing (WGS) technology, little was understood about resistome diversity that affects pathogen endemicity (15). Traditional typing methods generally exhibit low resolution, making monitoring and quantitatively comparing horizontal ARG transfer events challenging. The scarcity of appropriate models also poses a limitation. Salmonella, a pathogen with distinct geographical characteristics, has more than 90% of its serovars frequently categorized as geo-serotypes (16). Among the thousands of geo-serotypes, Salmonella enterica serovar Gallinarum (S. Gallinarum) is an avian-specific pathogen that causes severe mortality, with particularly detrimental effects on the poultry industry in low- and middle-income countries (17, 18). As once a globally prevalent pathogen in the 20th century, S. Gallinarum was listed by the World Organization for Animal Health (WOAH) and gradually became an endemic pathogen with sporadic outbreaks following the implementation of eradication programs in most high-income countries, making it a perfect model for study (19, 20).

Revised version: Before recent advancements in whole-genome sequencing (WGS) technology, little was understood about resistome diversity that affects pathogen endemicity (15). Traditional typing methods generally exhibit low resolution, making monitoring and quantitatively comparing horizontal ARG transfer events challenging. The scarcity of appropriate models also poses a limitation. Salmonella, a widespread foodborne and zoonotic pathogen with distinct geographical characteristics, has more than 90% of its serovars typically classified as geo-serotypes (16, 17). Among the thousands of geo-serotypes, Salmonella enterica serovar Gallinarum (S. Gallinarum) is an avian-specific pathogen that causes severe mortality, with particularly detrimental effects on the poultry industry in low– and middle-income countries (18–20). As once a globally prevalent pathogen in the 20th century, S. Gallinarum was listed by the World Organization for Animal Health (WOAH) and gradually became an endemic pathogen with sporadic outbreaks following the implementation of eradication programs in most high-income countries, making it a perfect model for study (21, 22).

Earlier version: Nowadays, antimicrobial therapy remains a priority choice against S. Gallinarum infections (21). The overuse or misuse of antimicrobials has led to increased epidemiological escalation of S. Gallinarum, with a regionally higher risk of AMR, especially for sulfonamides, penicillin, and tetracyclines (22, 23). To fill the gaps in understanding the evolution of S. Gallinarum under regional-associated AMR pressures and its adaptation to endemicity, we collected the most comprehensive set S. Gallinarum isolates, consisting of 580 genomes, spanning the period from 1920 to 2023. Using such unique WGS dataset, we investigated: 1) the population structure and potential geographical transmission history of S. Gallinarum at a single base level; 2) the dynamic resistome and mobilome changes that are associated with S. Gallinarum transition to an endemic variant; and 3) horizontal resistome transfer frequency and pattern between distinct regions.

Revised version: Nowadays, antimicrobial therapy remains a priority choice against S. Gallinarum infections (23). The overuse or misuse of antimicrobials has led to increased epidemiological escalation of S. Gallinarum, with a regionally higher risk of AMR, especially for sulfonamides, penicillin, and tetracyclines (24, 25). To fill the gaps in understanding the evolution of S. Gallinarum under regional-associated AMR pressures and its adaptation to endemicity, we collected the most comprehensive set S. Gallinarum isolates, consisting of 580 genomes, spanning the period from 1966 to 2023. Using such a unique WGS dataset, we investigated: 1) the population structure and potential geographical transmission history of S. Gallinarum at a single base level; 2) the dynamic resistome and mobilome changes that are associated with S. Gallinarum transition to an endemic variant; and 3) horizontal resistome transfer frequency and pattern between distinct regions.

Results

Global distribution of S. Gallinarum links with distinct sublineages

Earlier version: To understand the global geographic distribution and genetic relationships of S. Gallinarum, we assembled the most comprehensive S. Gallinarum WGS dataset (n=580), comprising 535 publicly available genomes and 45 newly sequenced genomes (24). The core-genome single nucleotide polymorphism sites (cgSNPs) were obtained by using the fully sequenced genome of S. Gallinarum R51 as a reference. Through hierarchical Bayesian analysis of cgSNPs, it was confirmed that S. Gallinarum divides into three biovars: S. Gallinarum biovar Pullorum (bvSP) (n=528/580, 91.03%), S. Gallinarum biovar Gallinarum (bvSG) (n=50/580, 8.6%), and S. Gallinarum biovar Duisburg (bvSD) (n=2/580, 0.34%). An association was identified between the biovar type of S. Gallinarum and its global geographic distribution. Most bvSP isolates were from Asia (436/528), with fewer occurrences in Europe (54/528), South America (12/528), and North America (5/528). In bvSG, South America (18/50, 36%) is the primary source of isolates, followed by North America (9/50, 18%), Europe (9/50, 18%), Africa (7/50, 14%), and Asia (3/50, 6%) (S1 Fig).

Revised version: To understand the global geographic distribution and genetic relationships of S. Gallinarum, we assembled the most comprehensive S. Gallinarum WGS dataset (n=580, Supplementary Table 1 and 2), comprising 535 publicly available genomes and 45 newly sequenced genomes (26). The core-genome single nucleotide polymorphism sites (cgSNPs) were obtained by using the fully sequenced genome of S. Gallinarum R51 as a reference. Through hierarchical Bayesian analysis of cgSNPs, it was confirmed that S. Gallinarum divides into three biovars: S. Gallinarum biovar Pullorum (bvSP) (n=528/580, 91.03%), S. Gallinarum biovar Gallinarum (bvSG) (n=50/580, 8.6%), and S. Gallinarum biovar Duisburg (bvSD) (n=2/580, 0.34%). An association was identified between the biovar type of S. Gallinarum and its global geographic distribution. Most bvSP isolates were from Asia (436/528), with fewer occurrences in Europe (54/528), South America (12/528), and North America (5/528). In bvSG, South America (18/50, 36%) is the primary source of isolates, followed by North America (9/50, 18%), Europe (9/50, 18%), Africa (7/50, 14%), and Asia (3/50, 6%) (S1 Fig).

Earlier version: From a lineage perspective, bvSP can be further classified into five lineages: L1 (10/528), L2a (45/528), L2b (163/528), L3b (169/528), and L3c (141/528). Importantly, L3a, previously considered distinct, has been revealed as a sub-lineage of L3c. The predominant lineage types vary across different continents. Regarding predominant bvSP sources, Asia stands out, with L3c, L3b, and L2b being the main lineage types. On the other hand, L3b and L2a are more widespread in Europe and America, respectively (Fig 1a). For the bvSP strains from Asia included in our dataset, we found that all originated from China. To further investigate the distribution of bvSP across different regions in China, we categorized them into three distinct regions: southern, eastern, and northern (Supplementary Table 3). Our findings indicate significant variations in bvSP across various regions from geo-temporal aspects. Prevalence was highest in the eastern region (233/436), with lower incidences in the northern (137/436) and southern regions (65/436). Interestingly, the predominant bvSP lineage type also varied. In eastern and southern China, L2b and L3c exhibited as the predominant lineage types, while in northern China, L3b and L3c held nearly equivalent positions (Fig 1c).

Revised version: From a lineage perspective, bvSP can be further classified into five lineages: L1 (10/528), L2a (45/528), L2b (163/528), L3b (169/528), and L3c (141/528). Importantly, L3a, previously considered distinct, has been revealed as a sub-lineage of L3c. The predominant lineage types vary across different continents. Regarding predominant bvSP sources, Asia stands out, with L3c, L3b, and L2b being the main lineage types. On the other hand, L3b and L2a are more widespread in Europe and America, respectively (Fig 1a). For the bvSP strains from Asia included in our dataset, we found that all originated from China. To further investigate the distribution of bvSP across different regions in China, we categorized them into three distinct regions: southern, eastern, and northern (Supplementary Table 3). Our findings indicate significant variations in bvSP across various regions from geo-temporal aspects. Prevalence was highest in the eastern region (233/436), with lower incidences in the northern (137/436) and southern regions (65/436). Interestingly, the predominant bvSP lineage type also varied. In eastern and southern China, L2b and L3c exhibited as the predominant lineage types, while in northern China, L3b and L3c held nearly equivalent positions (Fig 1b).

Genetic diversity of S. Gallinarum biovar Pullorum (bvSP) by geography and time.

Earlier version: Different colours were used to represent the various lineages of bvSP: fuchsia for L1, orange for L2a, pink for L2b, red for L3b, and green for L3c. (a) The bvSP in the dataset is classified into five continents based on their isolation regions. The bar graph illustrates the distribution of lineage-specific bvSP across continents, depicted as percentages. The total sample size for each bar is indicated at the top. (b) The bar graph shows the distribution of bvSP isolated in China by lineage per decade. The total sample size for each bar is also indicated on the right side. (c) Geographic distribution of bvSP isolated from China. Doughnut charts in the map show the proportion of lineage types of bvSP collected in the corresponding region, with the total number of isolates in the center.

Revised version: Different colours were used to represent the various lineages of bvSP: fuchsia for L1, orange for L2a, pink for L2b, red for L3b, and green for L3c. (a) The bvSP in the dataset is classified into five continents based on their isolation regions. The bar graph illustrates the distribution of lineage-specific bvSP across continents, depicted as percentages. The total sample size for each bar is indicated at the top. (b) Geographic distribution of bvSP isolated from China. Doughnut charts in the map show the proportion of lineage types of bvSP collected in the corresponding region, with the total number of isolates in the center. (c) The bar graph shows the distribution of bvSP isolated in China by lineage per decade. The total sample size for each bar is also indicated on the right side.

Earlier version: An analysis of the temporal prevalence of bvSP in China revealed a gradual replacement pattern of the lineages. Before the 2000s, the predominant lineage type was L3b. Then, from the 2000s to the 2010s, L3b and L1 were continually replaced by L3c, with L2b becoming the predominant lineage after the 2010s (Fig 1b). However, the replacement pattern varied among different regions. Specifically, the same pattern was observed in eastern China, but for northern and southern China, only replacements from L3b to L2b and L1 to L3c were observed (S2 Fig). These findings might indicate accelerated localized adaptation of bvSP.

Revised version: An analysis of the temporal prevalence of bvSP in China revealed a gradual replacement pattern of the lineages. Before the 2000s, the predominant lineage type was L3b. Then, from the 2000s to the 2010s, L3b and L1 were continually replaced by L3c, with L2b becoming the predominant lineage after the 2010s (Fig 1c). However, the replacement pattern varied among different regions. Specifically, the same pattern was observed in eastern China, but for northern and southern China, only replacements from L3b to L2b and L1 to L3c were observed (S2 Fig). These findings might indicate accelerated localized adaptation of bvSP.

Genomics portrait bvSP historic transmission

Considering the historical pandemic status of bvSP, we then investigated its global geographic transmission routes to understand its evolution into an endemic pathogen. L2b and L3b were identified as the dominant global lineages due to their antimicrobial resistance risks, and potential intercontinental transmission events. Other lineages were also analyzed using a Bayesian model with a relaxed molecular clock to infer their historical evolution (S3-S5 Figs). The temporal structure was verified (S6 Fig). Our findings show that the origin of L3b in China can be traced back to as early as 1683 (95% CI: 1608 to 1839). In contrast, the earliest possible origin of L2b in China dates back to 1880 (95% CI: 1838 to 1902) (Fig 2). Then, we specifically estimated the time points for the first intercontinental transmission events for the two major lineages. Our results indicate that L2b, likely underwent two major intercontinental transmission events. The first occurred around 1893 (95% CI: 1870 to 1918), with transmission from China to South America. The second major transmission event occurred in 1923 (95% CI: 1907 to 1940), involving the spread from South America to Europe. In contrast, the transmission pattern of L3b, an S. Gallinarum lineage originating in China, appears relatively more straightforward. It underwent only one intercontinental transmission event, from China to Europe, likely around 1790 (95% CI: 1661 to 1890) (S7 Fig).

Phylogenetic tree of S. Gallinarum L2b and L3b based on a spatiotemporal Bayesian framework.

The phylogenetic tree on the left was constructed using a reference-mapped multiple core-genome SNPs sequence alignment, with recombination regions detected and removed by Gubbins. The spatiotemporal Bayesian framework was configured with the “GTR” substitution model, 4 Gamma Category Count, “Relaxed Clock Log Normal” model, “Coalescent Bayesian Skyline” tree prior model, and a Markov Chain Monte Carlo (MCMC) chain length of 100,000,000, with sampling every 10,000 iterations. Convergence was assessed using Tracer, ensuring all parameter effective sampling sizes (ESS) exceeded 200. Evolutionary time is represented by the length of the branches. The heatmap on the right displays, respectively, the sequence type (ST), region of isolation, and the number (Num.) of antimicrobial resistance genes (ARGs) carried by the corresponding Salmonella Gallinarum.

Endemic isolates further support localized dissemination by bvSP

Earlier version: To investigate the dissemination pattern of bvSP, we obtained forty-five newly isolated bvSP from 734 samples (6.1% overall isolation rate) collected from diseased chickens at two farms in Yueqing and Taishun, Zhejiang Province. Those bvSP can be classified into two dominant sequence types (STs): ST3717 and ST92. Isolates from Yueqing were all identified as ST3717, while most Taishun isolates were characterized as ST92. Only the E404 isolate was identified as ST2151 (Fig 3a). Of note, a lineage-preferential association was also observed, with bvSP isolated from Taishun belonging to L3b and those from Yueqing belonging to L2b. To further understand the genetic relationship among isolates in the two regions, we calculated cgSNPs distances between isolates. Isolates with cgSNP distances less than five were determined as the threshold to be involved in the same transmission event. Interestingly, genetic clustering revealed that isolates from the same farm exhibited a significantly close genetic correlation (Fig 3b). No historical transmission events of bvSP were found between Taishun and Yueqing.

Revised version: To investigate the dissemination pattern of bvSP, we obtained forty-five newly isolated bvSP from 734 samples (6.1% overall isolation rate) collected from diseased chickens at two farms in Yueqing and Taishun, Zhejiang Province. Those bvSP can be classified into two dominant sequence types (STs): ST3717 and ST92. Isolates from Yueqing were all identified as ST3717, while most Taishun isolates were characterized as ST92. Only the E404 isolate was identified as ST2151 (Fig 3a). Of note, a lineage-preferential association was also observed, with bvSP isolated from Taishun belonging to L3b and those from Yueqing belonging to L2b. To further understand the genetic relationship among isolates in the two regions, we calculated cgSNPs distances between isolates. Isolates with cgSNP distances less than two were determined as the threshold to be involved in the same transmission event. Interestingly, genetic clustering revealed that isolates from the same farm exhibited a significantly close genetic correlation (Fig 3b). No historical transmission events of bvSP were found between Taishun and Yueqing.

Removed in the revised version: Genomic characteristics of 45 newly isolated S. Gallinarum biovar Pullorum (bvSP) strains from Taishun and Yueqing, Zhejiang Province, China.

Earlier version: (a) The phylogenetic tree, constructed using cgSNPs, categorizes the 45 bvSP strains into two distinct lineages (L3b and L2b). On the left side of the heatmap, information on isolation regions, sequence types (ST), and sampling times are displayed, with various colors indicating different categories as specified on the left. The right section of the heatmap presents a detailed matrix showing plasmids, ARGs, and cgSNP distances. The presence of a plasmid or ARG in an isolate is denoted by gray shading, while absence is indicated by white. cgSNPs below five were used as the threshold, with darker circles signifying a higher probability of transmission between isolate pairs. (b) The average cgSNP distance between isolates from Taishun and Yueqing. E404 led to an increase in the mean cgSNP distance of bvSP from Taishun. (c) Invasiveness index of bvSP in Taishun and Yueqing. The results show a higher invasiveness index for bvSP isolated from Taishun, indicating that bvSP isolated from Taishun might have greater invasive capabilities among vulnerable hens.

Revised version: Genomic characteristics of 45 newly isolated S. Gallinarum biovar Pullorum (bvSP) strains from Taishun and Yueqing, Zhejiang Province, China. (a) The phylogenetic tree, constructed using cgSNPs, categorizes the 45 bvSP strains into two distinct lineages (L3b and L2b). On the left side of the heatmap, information on isolation regions, sequence types (ST), and sampling times are displayed, with various colors indicating different categories as specified on the left. The right section of the heatmap presents a detailed matrix showing plasmids, ARGs, and cgSNP distances. The presence of a plasmid or ARG in an isolate is denoted by gray shading, while absence is indicated by white. cgSNPs below two were used as the threshold, with red circles signifying a higher probability of transmission between isolate pairs. (b) The average cgSNP distance between isolates from Taishun and Yueqing. E404 led to an increase in the mean cgSNP distance of bvSP from Taishun. (c) Invasiveness index of bvSP in Taishun and Yueqing. The results show a higher invasiveness index for bvSP isolated from Taishun, indicating that bvSP isolated from Taishun might have greater invasive capabilities among vulnerable hens.

Earlier version: Other genetic characteristics also exhibited a similar pattern. The ARG and plasmid profile disclosed that IncFII(S) and ColpVC plasmids exhibited the highest carriage percentages (100%) among the 45 bvSP isolates. The IncX1 plasmid was exclusively identified in the E404 isolate. Furthermore, E404 harboured two distinctive resistance genes, blaTEM-1B and sul2, in contrast to the absence of ARGs in the remaining bvSP isolates. We also estimated the invasiveness index to assess the invasiveness of isolated bvSP and showed that bvSP from the same farm had comparable invasive abilities. However, isolates from Taishun demonstrated a higher invasive ability than those from Yueqing (Fig 3c, Supplementary Table 6).

Revised version: Other genetic characteristics also exhibited a similar pattern. The ARG and plasmid profile disclosed that IncFII(S) and ColpVC plasmids exhibited the highest carriage percentages (100%) among the 45 bvSP isolates. The IncX1 plasmid was exclusively identified in the E404 isolate. Furthermore, E404 harboured two distinctive resistance genes, blaTEM-1B and sul2, in contrast to the absence of ARGs in the remaining bvSP isolates. We also estimated the invasiveness index to assess the invasiveness of isolated bvSP and showed that bvSP from the same farm had comparable invasive abilities. However, isolates from Taishun demonstrated a higher invasive ability than those from Yueqing (Fig 3c, Supplementary Table 4).

Earlier version: Furthermore, we simulated potential transmission events between the bvSP strains isolated from Zhejiang Province (n=95) and those from China with available provincial information (n=435) using five SNPs as the threshold. As a result, the transmission events showed a strong geographic-preferential distribution. We identified a total of 91 potential transmission events, all of which occurred exclusively within Zhejiang Province. No inter-provincial transmission events were detected, supporting the statement that bvSP in China is a highly localized pathogen (S8 Fig, Supplementary Table 8).

Revised version: Furthermore, we simulated potential transmission events between the bvSP strains isolated from Zhejiang Province (n=95) and those from China with available provincial information (n=435) using two SNPs as the threshold. As a result, the transmission events showed a strong geographic-preferential distribution. We identified a total of 53 potential transmission events, all of which occurred exclusively within Zhejiang Province. No inter-provincial transmission events were detected, supporting the statement that bvSP in China is a highly localized pathogen (S8 Fig, Supplementary Table 5).

Mobilome and resistome distinct lineage-preferential association in a regional manner

Earlier version: The mobilome drives partitions of the resistome and plays a pivotal role in shaping the ecological niches and local adaptation of bacteria. Therefore, a quantitative evaluation of the resistome and mobilome is necessary to enhance the understanding of the localized distribution of bvSP. For S. Gallinarum, a total of 13 classes of ARGs, which can be further classified into six categories, were identified. The results revealed that bvSP exhibited a significantly greater resistome than bvSG and bvSD (Fig 4a). Among them, sul2 (196/528, 37.1%) has the highest prevalence, followed by blaTEM-1B (183/528, 34.7%) and tet(A) (104/528, 19.7%). Moreover, the resistome also demonstrated a lineage-preferential association. We observed that L3b is more inclined to carry blaTEM-1B and sul2 than other lineages, whereas tet(A) exclusively exists within L3 (Supplementary Table 4). The diversity of resistome carried by L3 may elevate the risk of multi-drug resistance (MDR). Notably, AMR risks vary among different regions of China, with the highest risk observed in southern China, followed by the northern and eastern regions (Fig 4c). Interestingly, tet(A), blaTEM-1B, and sul2 were predominant resistome types prevalent across all regions of China. Meanwhile, aadA5 was typically more commonly found in southern China (Fig 4b).

Revised version: The mobilome drives partitions of the resistome and plays a pivotal role in shaping the ecological niches and local adaptation of bacteria. Therefore, a quantitative evaluation of the resistome and mobilome is necessary to enhance the understanding of the localized distribution of bvSP. For S. Gallinarum, a total of 13 classes of ARGs, which can be further classified into six categories, were identified. The results revealed that bvSP exhibited a significantly greater resistome than bvSG and bvSD (Fig 4a). Among them, sul2 (196/528, 37.1%) has the highest prevalence, followed by blaTEM-1B (183/528, 34.7%) and tet(A) (104/528, 19.7%). Moreover, the resistome also demonstrated a lineage-preferential association. We observed that L3b is more inclined to carry blaTEM-1B and sul2 than other lineages, whereas tet(A) exclusively exists within L3 (Supplementary Table 6). The diversity of resistome carried by L3 may elevate the risk of multi-drug resistance (MDR). Notably, AMR risks vary among different regions of China, with the highest risk observed in southern China, followed by the northern and eastern regions. Interestingly, tet(A), blaTEM-1B, and sul2 were predominant resistome types prevalent across all regions of China. Meanwhile, aadA5 was typically more commonly found in southern China (Fig 4b and c).

The ARGs carried by S. Gallinarum.

(a) The phylogenetic tree was constructed using cgSNPs, revealing three S. Gallinarum biovars: S. Gallinarum biovar Pullorum (bvSP), S. Gallinarum biovar Duisburg (bvSD), and S. Gallinarum biovar Gallinarum (bvSG). Additionally, Salmonella serovar Enteritidis (SE.) is represented by a gray line. Further, bvSP can be subdivided into five lineages: fuchsia (L1), orange (L2a), pink (L2b), red (L3b), and green (L3c). The heatmap on the right indicates the resistome carried by the corresponding Salmonella. (b) The dominant resistome types in different regions of China. The y-axis represents the percentage of each dominant resistome. The total sample size for each bar is indicated at the top. (c) The average number of resistome carried by bvSP is from different regions of China.

Earlier version: Four categories of the mobilome—prophage, plasmid, transposon, and integron—were characterized, revealing strong lineage-specific patterns (S9a Fig, Supplementary Table 5). In silico analyses revealed prophages emerge as the most diverse mobilome type, with Entero_mEp237, Escher_500465_1, and Klebsi_4LV2017 being more likely associated with L1, L2a, and L2b, respectively, while the carrier percentage of Escher_pro483 is significantly higher among L3c. Interestingly, Salmon_SJ46, Gifsy_2, Escher_500465_2, and Shigel_SfIV were consistently observed across all bvSP lineages and were found to be duplicated on the chromosome. For plasmids, IncFII(S) was predominantly carried across all bvSP lineages (520/528, 98.5%) (S9b Fig). Notably, ColpVC had the lowest carriage percentage among L1 (2/10, 20%), whereas IncX1 was explicitly associated with L3c (140/141, 99.3%).

Revised version: Four categories of the mobilome—prophage, plasmid, transposon, and integron—were characterized, revealing strong lineage-specific patterns (S9a Fig, Supplementary Table 7). In silico analyses revealed prophages emerge as the most diverse mobilome type, with Entero_mEp237, Escher_500465_1, and Klebsi_4LV2017 being more likely associated with L1, L2a, and L2b, respectively, while the carrier percentage of Escher_pro483 is significantly higher among L3c. Interestingly, Salmon_SJ46, Gifsy_2, Escher_500465_2, and Shigel_SfIV were consistently observed across all bvSP lineages and were found to be duplicated on the chromosome. For plasmids, IncFII(S) was predominantly carried across all bvSP lineages (520/528, 98.5%) (S9b Fig). Notably, ColpVC had the lowest carriage percentage among L1 (2/10, 20%), whereas IncX1 was explicitly associated with L3c (140/141, 99.3%).

The diversity of transposons and integrons is comparatively lower in bvSP. Specifically, we identified seven types of transposons and four types of integrons. Transposons were found to be more abundant, comprising a total of 255 in 528 bvSP. Among them, Tn801 (185/528, 35%) had the highest carriage rate, followed by Tn1721 (48/528, 9.1%). Interestingly, L3c is more likely to carry Tn801 and Tn1721 transposons than other lineages. In contrast, a total of 42 integrons were identified among the 528 bvSP, all of which belong to class I: In498 (n=22), In1440 (n=14), In473 (n=4), and In282 (n=2). The integrons also exhibited a lineage-preferential distribution, predominantly carried by L3b.

From a geographic distribution perspective, bvSP from northern and southern China carried the most extensive mobilome, followed by those from eastern China. Interestingly, bvSP isolated from northern China typically have more diverse mobilome types (S10 Fig).

Plasmid and transposon guide resistome geo-temporal dissemination

Variations in regional antimicrobial use may result in uneven pressure for selecting AMR. The mobilome is considered the primary reservoir for spreading resistome, and a consistent trend between the resistome and the mobilome has been observed across different lineages, from L1-L3c. We observed an overall gradual rise in the resistome quantity carried by bvSP across various lineages, correlating with the total mobilome content (S11 Fig). Furthermore, we investigated the interplay between particular mobile elements and resistome types in bvSP.

In bvSP, the predominant resistome is mainly associated with certain types of mobilome (Fig 5a). Specifically, regarding blaTEM_1B, tet(A), and sul1, which are highly prevalent among bvSP, we found most blaTEM_1B were carried by the transposon Tn801 and Tn1721, the plasmid IncX1, and the prophage SJ46. Additionally, tet(A) was primarily associated with the transposon Tn1721 and the plasmid IncX1. Integrons also facilitated the dissemination of the resistome. For instance, In498 and In1440 facilitate the dissemination of sul1, leading to quinolone resistance. Our results indicate that plasmids and transposons are the predominant reservoirs for the resistome in bvSP, with prophages and integrons following closely behind (Fig 5b). Interestingly, we observed that the primary reservoir for the resistome shows regional variations. In the southern and northern regions of China, the primary reservoirs for blaTEM_1B are Tn801, Tn1721, IncX1, and prophage SJ46. However, Tn1721 was less frequently found in bvSP, which carries blaTEM-1B from the eastern region. Similarly, the reservoir for tet(A) exhibits greater diversity in the east region but less diversity in the southern region (Fig 5c-e).

The primary source of resistome is carried by distinct mobilome.

Different font colours denote various mobilome types. Specifically, orange signifies Integrons, red denotes transposons, blue represents plasmids, and green indicates prophages. A black font is utilized to distinguish the categories of resistome. The connecting line between the resistome and the mobilome represents the potential carrying relationship. (a) The mobilome-carried ARGs among S. Gallinarum biovar Pullorum (bvSP). (b) The average number of ARGs carried by the four MGEs in bvSP. The unpaired t-test was used to compare the differences between isolates, with P < 0.05 considered statistically significant. (c-e) The mobilome-carried ARGs among bvSP isolated from the eastern, southern, and northern regions of China.

Horizontal transfer of resistome occurs widely in localized bvSP

Earlier version: Horizontal transfer of the resistome facilitates the acquisition of AMR among bacteria, which may record the distinct acquisition event in the bacterial genome. To compare these events in a geographic manner, we further investigated the HGT frequency of each ARG carried by bvSP isolated from China and explored the HGT frequency of resistome between three defined regions. Potentially horizontally transferred ARGs were defined as those with perfect identity (100% identity and 100% coverage) and were located on MGEs across different strains (Fig. 6a). We first categorized a total of 621 ARGs carried by 436 bvSP isolates collected in China and found that 415 ARGs were located on MGEs (Supplementary Table 9). After excluding the ARGs not associated with MGEs, our findings reveal that horizontal gene transfer of ARGs is widespread among Chinese bvSP isolates, with an overall transfer rate of 92%. Specifically, 50% of the ARGs exhibited an HGT frequency of 100%, indicating that these ARGs might underwent extensive frequent horizontal transfer events (Fig. 6b). It is noteworthy that certain resistance genes, such as tet(A), aph(3’’)-Ib, and aph(6)-Id, appear to be less susceptible to horizontal transfer.

Revised version: Horizontal transfer of the resistome facilitates the acquisition of AMR among bacteria, which may record the distinct acquisition event in the bacterial genome. To compare these events in a geographic manner, we further investigated the HGT frequency of each ARG carried by bvSP isolated from China and explored the HGT frequency of resistome between three defined regions. Potentially horizontally transferred ARGs were defined as those with perfect identity (100% identity and 100% coverage) and were located on MGEs across different strains (Fig. 6a). We first categorized a total of 621 ARGs carried by 436 bvSP isolates collected in China and found that 415 ARGs were located on MGEs (Supplementary Table 8). After excluding the ARGs not associated with MGEs, our findings reveal that horizontal gene transfer of ARGs is widespread among Chinese bvSP isolates, with an overall transfer rate of 92%. Specifically, 50% of the ARGs exhibited an HGT frequency of 100%. Using the HGTphyloDetect pipeline, we verified the HGT transmission potential for each ARG sequence. The results demonstrated that ARGs with an HGT frequency greater than zero (blaTEM-1B, sul1, dfrA17, aadA5, sul2, aph(3’’)-Ib, tet(A), aph(6)-Id) all had Alien Index scores exceeding 45 and out_perc values greater than 90, indicating that these ARGs may have undergone extensive horizontal transfer events (Fig. 6b, Supplementary Table 9). It is noteworthy that certain resistance genes, such as tet(A), aph(3’’)-Ib, and aph(6)-Id, appear to be less susceptible to horizontal transfer.

The HGT frequency of the resistome among S. Gallinarum biovar Pullorum (bvSP) isolated from China.

Earlier version: (a) Workflow for identification of horizontally transferred ARGs in Salmonella. (b) The x-axis represents the resistome of bvSP, while the y-axis represents the corresponding levels of HGT frequency. (c) The overall resistome HGT frequency level of bvSP isolated from the various regions of China. (d) The HGT frequency level of specific ARGs carried by bvSP isolated from various regions of China. Deeper colours mean higher HGT frequency. (e) The frequency of horizontal retransmission of the resistome between different regions of China. A higher value indicates more frequent transfer events of resistome between two regions. The colors on the map represent the number of bvSP isolates in each region. Darker colors indicate a higher number of bvSP isolates in that area.

Revised version: (a) Workflow for identification of horizontally transferred ARGs in Salmonella. (b) The x-axis represents the resistome of bvSP, while the y-axis represents the corresponding levels of HGT frequency. (c) The HGT frequency level of specific ARGs carried by bvSP isolated from various regions of China. Deeper colours mean higher HGT frequency. (d) The frequency of horizontal retransmission of the resistome between different regions of China. A higher value indicates more frequent transfer events of resistome between two regions. The colors on the map represent the number of bvSP isolates in each region. Darker colors indicate a higher number of bvSP isolates in that area.

Earlier version: However, different regions generally exhibited a considerable difference in resistome HGT frequency. Overall, bvSP from the southern areas in China showed the highest HGT frequency (HGT frequency=95%). The HGT frequencies for bvSP within the eastern and northern regions of China are lower, at 92% and 91%, respectively (Fig. 6c). For specifical ARG type, we found tet(A) is more prone to horizontal transfer in the southern region, and this proportion was considerably lower in the eastern region. Interestingly, certain ARGs such as aph(6)-Id, undergo horizontal transfer only within the eastern and northern regions of China (Fig. 6d). Notably, as a localized transmission pathogen, resistome carried by bvSP exhibited a dynamic potential among inter-regional and local demographic transmission, especially from northern region to southern region (HGT frequency=93%) (Fig. 6e, Supplementary Table 7).

Revised version: However, different regions generally exhibited a considerable difference in resistome HGT frequency. For specifical ARG type, we found tet(A) is more prone to horizontal transfer in the southern region, and this proportion was considerably lower in the eastern region. Interestingly, certain ARGs, such as aph(6)-Id, undergo horizontal transfer only within the eastern and northern regions of China (Fig. 6c). Notably, as a localized transmission pathogen, resistome carried by bvSP exhibited a dynamic potential among inter-regional and local demographic transmission, especially from northern region to southern region (HGT frequency=93%) (Fig. 6d, Supplementary Table 10).

Discussion

Earlier version: Over the past decades, the emergence of AMR has garnered significant attention from both public health agencies and the poultry industry (25). Region-associated pressures shaped the dissemination patterns of bacterial pathogens, transitioning from pandemic to endemic, and led to dynamic changes in the resistome. Here, utilizing S. Gallinarum as a prototype for an endemic pathogen that once spread globally, and assembling the most comprehensive S. Gallinarum WGS database, we have discovered that region-associated pressures can influence the AMR risks of S. Gallinarum within lineage- and region-specific patterns. These pressures might also influence lineage-distinct evolutionary structures and transmission histories of S. Gallinarum.

Revised version: Over the past decades, the emergence of AMR has garnered significant attention from both public health agencies and the poultry industry (27). Region-associated pressures shaped the dissemination patterns of bacterial pathogens, transitioning from pandemic to endemic, and led to dynamic changes in the resistome. Here, utilizing S. Gallinarum as a prototype for an endemic pathogen that once spread globally, and assembling the most comprehensive S. Gallinarum WGS database, we have discovered that region-associated pressures can influence the AMR risks of S. Gallinarum within lineage– and region-specific patterns. These pressures might also influence lineage-distinct evolutionary structures and transmission histories of S. Gallinarum.

Earlier version: The biovar types of S. Gallinarum have been well-defined as bvSP, bvSG, and bvSD historically (26). Among these, bvSP can be further subdivided into five lineages (L1, L2a, L2b, L3b, and L3c) using hierarchical Bayesian analysis. Different sublineages exhibited preferential geographic distribution, with L2b and L3b of bvSP being predominant global lineage types with a high risk of AMR. The historical geographical transmission was verified using a spatiotemporal Bayesian framework. The result shows that L3b was initially spread from China to Europe in the 18th-19th century, which may be associated with the European hen fever event in the mid-19th century (27). L2b, on the other hand, appears to have spread to Europe via South America, potentially contributing to the prevalence of bvSP in the United States. Considering the economic losses caused by bvSP, the United States, Europe, and other industrialized countries implemented eradication programs in the mid-19th century, thus eliminating the risk of pullorum disease. Implementing similar measures is challenging due to China’s vast geographical area and various economic factors, resulting in bvSP becoming an endemic pathogen, mainly in China.

Revised version: The biovar types of S. Gallinarum have been well-defined as bvSP, bvSG, and bvSD historically (28). Among these, bvSP can be further subdivided into five lineages (L1, L2a, L2b, L3b, and L3c) using hierarchical Bayesian analysis. Different sublineages exhibited preferential geographic distribution, with L2b and L3b of bvSP being predominant global lineage types with a high risk of AMR. The historical geographical transmission was verified using a spatiotemporal Bayesian framework. The result shows that L3b was initially spread from China to Europe in the 18th-19th century, which may be associated with the European hen fever event in the mid-19th century (29). L2b, on the other hand, appears to have spread to Europe via South America, potentially contributing to the prevalence of bvSP in the United States. Considering the economic losses caused by bvSP, the United States, Europe, and other industrialized countries implemented eradication programs in the mid-19th century, thus eliminating the risk of pullorum disease. Implementing similar measures is challenging due to China’s vast geographical area and various economic factors, resulting in bvSP becoming an endemic pathogen, mainly in China.

Added in the revised version: Additionally, forty-five new bvSP isolates were obtained from two different locations (Taishun and Yueqing, Zhejiang Province). The cgSNP distances (<2 SNPs) revealed potential transmission events occurring exclusively among isolates from the same location. Traceback analysis using isolates from Zhejiang Province and other Chinese isolates with available provincial information further supports that bvSP in China is a highly localized pathogen. Moreover, the invasiveness index of bvSP from Taishun and Yueqing suggests that different lineages of S. Gallinarum recovered from distinct regions may exhibit biological differences. Previous studies have shown that strains with higher invasiveness indexes tend to be more virulent in hosts (30, 31), potentially causing neurological or arthritic symptoms in S. Gallinarum infections. Furthermore, strains with varying invasiveness indexes have been confirmed to differ in their biofilm formation abilities and metabolic capacities for carbon compounds (32).

Removed in the revised version: Additionally, forty-five bvSP was isolated from Taishun and Yueqing, to explore the transmission pattern of bvSP in China. The cgSNP distances between bvSP isolates from Taishun and Yueqing, showed significant diversity, with cgSNP distances of less than 100 isolate pairs observed exclusively within a single region. The transmission history of 45 global bvSP isolates revealed that most transmission events occurred within a single region. Additionally, the number of transmission events decreased rapidly with increasing distance, confirming that bvSP in China belongs to a highly localized pathogen.

Removed in the revised version: Humanity appeared powerless against epidemics until the advent of antimicrobials, which are still the primary agents used to combat S. Gallinarum infection (28, 29). As a manifestation of localization-related forces, the emergence of AMR allows pathogens to acquire additional fitness benefits (30).

Earlier version: Our findings further indicate that the mobilome facilitates the acquisition of AMR among bvSP, mainly through plasmids and transposons, such as IncX1, IncQ1, and IncN type plasmids, as well as Tn801, Tn6205, and Tn721 type transposons. A survey conducted in China between 1962 and 2010 documented the association of class 1 integrons with AMR among bvSP (31). Our study further demonstrates that In498, In1440, In473, and In282, all belonging to class 1 integrons, play a significant role in carrying ARGs associated with sulfonamides, and trimethoprim. However, compared to other Salmonella serovars (32–34), the level of resistome in bvSP remains lower.

Revised version: Humanity appeared powerless against epidemics until the advent of antimicrobials, which are still the primary agents used to combat S. Gallinarum infection (33, 34). As a manifestation of localization-related forces, the emergence of AMR allows pathogens to acquire additional fitness benefits (35). Our findings further indicate that the mobilome facilitates the acquisition of AMR among bvSP, mainly through plasmids and transposons, such as IncX1, IncQ1, and IncN type plasmids, as well as Tn801, Tn6205, and Tn721 type transposons. A survey conducted in China between 1962 and 2010 documented the association of class 1 integrons with AMR among bvSP (36). Our study further demonstrates that In498, In1440, In473, and In282, all belonging to class 1 integrons, play a significant role in carrying ARGs associated with sulfonamides, and trimethoprim. However, compared to other Salmonella serovars (37–39), the level of resistome in bvSP remains lower.

Earlier version: Considering lineage prospects, recently emerged lineages present a higher risk of AMR, with a lineage-specific distribution of resistome. In L2b and L3c, the most prevalent groups of ARGs include blaTEM-1b and sul2. L3c, which exhibits the highest resistome richness with sul1, tet(A), aadA5, and dfrA17, has a higher likelihood of MDR. However, the transmission routes of resistome showed a high degree of consistency across continents. Globally, a small proportion of resistome carriers belonging to L2b and L3b were observed in the Americas and Europe, while isolates from Asia exhibited a notable rise in resistome carriage. Therefore, we speculate that the mobilome in bvSP may have been acquired locally in Asia, possibly due to increased antimicrobial usage and an intensive industrialized poultry farming pattern within the regional poultry industry (35). The L3c, being a local Asian lineage and more prone to acquiring additional resistome locally, provided further support.

Revised version: Considering lineage prospects, recently emerged lineages present a higher risk of AMR, with a lineage-specific distribution of resistome. In L2b and L3c, the most prevalent groups of ARGs include blaTEM-1b and sul2. L3c, which exhibits the highest resistome richness with sul1, tet(A), aadA5, and dfrA17, has a higher likelihood of MDR. However, the transmission routes of resistome showed a high degree of consistency across continents. Globally, a small proportion of resistome carriers belonging to L2b and L3b were observed in the Americas and Europe, while isolates from Asia exhibited a notable rise in resistome carriage. Therefore, we speculate that the mobilome in bvSP may have been acquired locally in Asia, possibly due to increased antimicrobial usage and an intensive industrialized poultry farming pattern within the regional poultry industry (40). The L3c, being a local Asian lineage and more prone to acquiring additional resistome locally, provided further support.

Earlier version: In China, the endemic region of Asia, an unexpected rising trend has been observed from the 1970s to the 2020s. However, the risk of AMR varies among distinct regions (36–39). Overall, the average number of ARGs carried by bvSP isolated from eastern China was the lowest, while it was higher among bvSP isolated from northern and southern China. This variation may be influenced by differences in the economy, climate, or farming practices among different regions. Interestingly, the widespread horizontal transfer of the resistome across regions seems to indicate that the resistome has a different mode of transmission compared to its pathogenic bacteria. Specifically, as a localized pathogen, our results show a low frequency of cross-regional transmission of bvSP (as shown in S5 Fig), but up to 90% of ARGs are transmitted cross-regionally. This risk has not been previously observed in S. Gallinarum, suggesting that potential intermediate hosts may play a key role.

Revised version: In China, the endemic region of Asia, an unexpected rising trend has been observed from the 1970s to the 2020s. However, the risk of AMR varies among distinct regions (41–44). Overall, the average number of ARGs carried by bvSP isolated from eastern China was the lowest, while it was higher among bvSP isolated from northern and southern China. This variation may be influenced by differences in the economy, climate, or farming practices among different regions. Interestingly, the widespread horizontal transfer of the resistome across regions indicates that the resistome has a different mode of transmission compared to its pathogenic bacteria. Specifically, as a localized pathogen, our results show a low frequency of cross-regional transmission of bvSP, but up to 90% of ARGs are transmitted cross-regionally. This risk has not been previously observed in S. Gallinarum, suggesting that potential intermediate hosts may play a key role.

Earlier version: In summary, the findings of this study highlight that S. Gallinarum remains a significant concern in developing countries, particularly in China. Compared to other regions, S. Gallinarum in China poses a notably higher risk of AMR, necessitating the development of additional therapies, i.e. vaccine, probiotics, bacteriophage therapy in response to the government’s policy aimed at reducing antimicrobial use (40). Furthermore, given the dynamic nature of S. Gallinarum risks across different regions, it is crucial to prioritize continuous monitoring in key areas, particularly in China’s southern regions where the extensive poultry farming is located. Lastly, from a One-Health perspective, controlling AMR in S. Gallinarum should not solely focus on local farming environments, with improved overall welfare on poultry and farming style. The breeding pyramid of industrialized poultry production should be targeted on the top, with enhanced and accurate detection techniques (41). More importantly, comprehensive efforts should be made to reduce antimicrobial usage overall and mitigate potential AMR transmission from environmental sources or other hosts (42–44).

Revised version: In summary, the findings of this study highlight that S. Gallinarum remains a significant concern in developing countries, particularly in China. Compared to other regions, S. Gallinarum in China poses a notably higher risk of AMR, necessitating the development of additional therapies, i.e. vaccine, probiotics, bacteriophage therapy in response to the government’s policy aimed at reducing antimicrobial use (45, 46). Furthermore, given the dynamic nature of S. Gallinarum risks across different regions, it is crucial to prioritize continuous monitoring in key areas, particularly in China’s southern regions where extensive poultry farming is located. Lastly, from a One-Health perspective, controlling AMR in S. Gallinarum should not solely focus on local farming environments, with improved overall welfare on poultry and farming style. The breeding pyramid of industrialized poultry production should be targeted at the top, with enhanced and accurate detection techniques (47). More importantly, comprehensive efforts should be made to reduce antimicrobial usage overall and mitigate potential AMR transmission from environmental sources or other hosts (48–50).

However, the current study has some limitations. Firstly, despite assembling the most comprehensive WGS database for S. Gallinarum from public and laboratory sources, there are still biases in the examined collection. The majority (438/580) of S. Gallinarum samples were collected from China, possibly since the WGS is a technology that only became widely available in the 21st century. This makes it impractical to sequence it on a large scale in the 20th century, when S. Gallinarum caused a global pandemic. So, we suspect that human intervention in the development of this epidemic is the main driving force behind the fact that most of the strains in the data set originated in China. In our future work, we aim to actively gather more data to minimize potential biases within our dataset, thereby improving the robustness and generalizability of our findings. Secondly, in silico analysis relies heavily on sophisticated, continuously updated databases. Although we utilized the most up-to-date databases, the exact prevalence of mobilome genetic elements (MGEs) and ARGs may be underestimated. Nevertheless, by using the most extensive global WGS data from S. Gallinarum, we elucidate a dynamic resistome evolution framework in a single bacterial pathogen from pandemic to endemic, which is hallmarked by a gradual increase in antimicrobial resistance risks and a stepwise resistome adaptation. The insights gained in this study will be helpful in further prevention and surveillance of the AMR risks of localized pathogens.

Material and methods

Ethics statements

The protocols used for animal studies were approved by the Committee of the Laboratory Animal Center of Zhejiang University (ZJU20190094; ZJU20220295).

Bacterial isolates

All 734 samples of dead chicken embryos aged 19 to 20 days were collected from Taishun and Yueqing in Zhejiang Province, China. After a thorough autopsy, the liver, intestines, and spleen were extracted and added separately into 2 mL centrifuge tubes containing 1 mL PBS. The organs were then homogenized by grinding. In the initial enrichment phase, we utilized Buffered Peptone Water (BPW, Haibo Biotechnology Co, Qingdao, China), employing a 1:9 dilution ratio (sample in PBS: BPW). Subsequently, the composite was incubated at 37°C for 16-18 hours in a rotary incubator set to 180 rpm. For selective enrichment, Tetrathionate Broth Base (TTB, Land Bridge Biotechnology Co, Beijing, China), fortified with iodine solution and brilliant green solution (both from Land Bridge Biotechnology Co, Beijing, China), was employed at a ratio of 1:10 (sample in BPW: TTB). The mixture was subjected to incubation at 42°C for a duration ranging between 22 and 26 hours within a rotary incubator set at 180 rpm. Isolated Salmonella colonies from positive samples were obtained by sub-culturing selectively enriched samples on Xylose Lysine Deoxycholate (XLD, Land Bridge Technology Co, Beijing, China) agar, followed by an 18–22-hour incubation at 37°C. Typical and pure colonies were selected after sub-culturing on XLD agar and transferred into Luria-Bertani (LB) broth. Finally, the transferred bacterial culture underwent an additional incubation for 18-22 hours at 37°C in a rotary incubator operating at 180 rpm.

DNA extraction and genomic assembly

Earlier version: The DNA of forty-five isolates was extracted using the Vazyme Fastpure® Bacteria DNA Isolation Mini Kit (Vazyme Biotech Co., Ltd.), which was then quantified using the NanoDrop1000 system (Thermo Fisher Scientific, USA). Subsequently, DNA libraries were constructed and subjected to sequencing using the Illumina Novaseq 6000 platform (Beijing Novogene Co., Ltd). Assembly of genome sequences was performed by SPAdes (45) v3.12.0 with default parameters.

Revised version: The DNA of forty-five isolates was extracted using the Vazyme Fastpure® Bacteria DNA Isolation Mini Kit (Vazyme Biotech Co., Ltd.), which was then quantified using the NanoDrop1000 system (Thermo Fisher Scientific, USA). Subsequently, DNA libraries were constructed and subjected to sequencing using the Illumina Novaseq 6000 platform (Beijing Novogene Co., Ltd). Assembly of genome sequences was performed by SPAdes (51) v3.12.0 with default parameters.

Global dataset assembly

To provide global context, we expanded the dataset by including 540 additional WGS genome data, consisting of bvSP (n=483), bvSG (n=50), bvSD (n=2), and Salmonella enterica serovar Enteritidis (n=5). The inclusion of five strains of Salmonella enterica serovar Enteritidis enhanced evolutionary analyses. Notably, in this collection of 540 genomic data from public databases, 325 sequences were previously preserved in our laboratory.

Earlier version: All WGS data passed strict quality control according to criteria set by the European Reference Laboratory (46). Genomic data exceeding 500 contigs and having an N50 of less than 30,000 were excluded. Lastly, the bacterial species associated with each genomic data set was confirmed by utilizing KmerFinder (47) v3.2. Finally, the assembled database includes a total of 585 high-quality sequences.

Revised version: All WGS data passed strict quality control according to criteria set by the European Reference Laboratory (52). Genomic data exceeding 500 contigs and having an N50 of less than 30,000 were excluded. Lastly, the bacterial species associated with each genomic data set was confirmed by utilizing KmerFinder (53) v3.2. Finally, the assembled database includes a total of 585 high-quality sequences.

Genotyping and phylogenetic analyses

Earlier version: Snippy v.4.4.5 was used to identify cgSNPs, with the complete genome of R51 used as a reference. The optimal model (TVM+F) was determined and employed to construct the maximum-likelihood phylogenetic tree using IQ-TREE v.1.6 (48). The final phylogenetic tree was annotated and visualized using iTOL v.6.0 (49).

Revised version: Snippy v.4.4.5 was used to identify cgSNPs, with the complete genome of R51 used as a reference. The optimal model (TVM+F) was determined and employed to construct the maximum-likelihood phylogenetic tree using IQ-TREE v.1.6 (54). The final phylogenetic tree was annotated and visualized using iTOL v.6.0 (55).

Earlier version: The global assembly of S. Gallinarum underwent a population structure analysis using the RhierBAPS v1.1.4 (50), based on the Hierarchical Bayesian clustering (BHC) algorithm and employing cgSNPs as input. The calculation of population structure utilized R package v.4.3.1, supplemented by the R packages “phytools” v.2.0.3, “ape” v.5.7.1, and “rhierbaps” v.1.1.4. Default parameters were employed for all population structure analysis.

Revised version: The global assembly of S. Gallinarum underwent a population structure analysis using the RhierBAPS v1.1.4 (56), based on the Hierarchical Bayesian clustering (BHC) algorithm and employing cgSNPs as input. The calculation of population structure utilized R package v.4.3.1, supplemented by the R packages “phytools” v.2.0.3, “ape” v.5.7.1, and “rhierbaps” v.1.1.4. Default parameters were employed for all population structure analysis.

Temporal and phylogeographical analysis

Earlier version: To examine the emergence and geographical transfers, we first utilized Gubbins v.2.3 (51) to eliminate the recombination regions for each lineags (S12 Fig). Then, TreeTime (52) was used to assess the temporal structure. This was accomplished through regression analysis of the root-to-tip branch distances within the maximum likelihood tree, considering the sampling date as a variable. Then, the strains within the dataset underwent phylogeographical reconstruction via Bayesian Evolutionary Analysis of Sampled Trees (BEAST) (53) v.2.5. To determine the optimal model for running BEAST, we tested a total of six combinations in the initial phase of our study. These combinations included the strict clock, relaxed lognormal clock, and three population models (Bayesian SkyGrid, Bayesian Skyline, and Constant Size). Before conducting the full BEAST analysis, we evaluated each combination using a Markov Chain Monte Carlo (MCMC) analysis with a total chain length of 100 million and sampling every 10,000 iterations. We then summarized the results using NSLogAnalyser and determined the optimal model based on the marginal likelihood value for each combination. The results indicated that the model incorporating the Bayesian Skyline, and the relaxed lognormal clock yielded the highest marginal likelihood value in our sample. Consequently, we proceeded to perform a time-calibrated Bayesian phylogenetic inference analysis for each lineage. The following settings were configured: the “GTR” substitution model, “4 gamma categories”, the “Relaxed Clock Log Normal” model, the “Coalescent Bayesian Skyline” tree prior. Convergence was assessed using Tracer, with all parameter effective sampling sizes (ESS) exceeding 200. Maximum clade credibility trees were generated using TreeAnnotator v.2.6.7. Finally, key divergence time points (with 95% credible intervals) were estimated and the phylogenetic tree was visualized using Figtree v.1.4.3.

Revised version: To examine the emergence and geographical transfers, we first utilized Gubbins v.2.3 (57) to eliminate the recombination regions for each lineags (S12 Fig). Then, TreeTime (58) was used to assess the temporal structure. This was accomplished through regression analysis of the root-to-tip branch distances within the maximum likelihood tree, considering the sampling date as a variable. Then, the strains within the dataset underwent phylogeographical reconstruction via Bayesian Evolutionary Analysis of Sampled Trees (BEAST) (59) v.2.5. To determine the optimal model for running BEAST, we tested a total of six combinations in the initial phase of our study. These combinations included the strict clock, relaxed lognormal clock, and three population models (Bayesian SkyGrid, Bayesian Skyline, and Constant Size). Before conducting the full BEAST analysis, we evaluated each combination using a Markov Chain Monte Carlo (MCMC) analysis with a total chain length of 100 million and sampling every 10,000 iterations. We then summarized the results using NSLogAnalyser and determined the optimal model based on the marginal likelihood value for each combination. The results indicated that the model incorporating the Bayesian Skyline, and the relaxed lognormal clock yielded the highest marginal likelihood value in our sample. Consequently, we proceeded to perform a time-calibrated Bayesian phylogenetic inference analysis for each lineage. The following settings were configured: the “GTR” substitution model, “4 gamma categories”, the “Relaxed Clock Log Normal” model, the “Coalescent Bayesian Skyline” tree prior. Convergence was assessed using Tracer, with all parameter effective sampling sizes (ESS) exceeding 200. Maximum clade credibility trees were generated using TreeAnnotator v.2.6.7. Finally, key divergence time points (with 95% credible intervals) were estimated and the phylogenetic tree was visualized using Figtree v.1.4.3.

SNP distance-based S. Gallinarum geographical tracing

Earlier version: To understand potential transmission events, we calculated the cgSNPs distance between distinct bacterial strains employing the SNP-dists. We estimated the overall evolutionary rate of the S. Gallinarum using BEAST. We applied the methodology described previously (54). The numbers of SNPs per year were determined by multiplying the evolutionary rates estimated with BEAST by the number of core SNP sites identified in the alignments. We hypothesize that a slower evolutionary rate in bacteria typically requires a lower SNP threshold when tracing transmission events using SNP distance analysis. Previous research found an average evolutionary rate of 1.97 SNPs per year (95% HPD, 0.48 to 4.61) across 22 different Salmonella serotypes. Our updated BEAST estimation for the evolutionary rate of S. Gallinarum suggests it is approximately 0.74 SNPs per year (95% HPD, 0.42 to 1.06). Based on these findings, as well as our previous experience with similar studies (55), we set a threshold of 5 SNPs. Additionally, the geographic location was considered to enhance the precision of speculation regarding potential transmission sites and transmission direction.

Revised version: To understand potential transmission events, we calculated the cgSNPs distance between distinct bacterial strains employing the SNP-dists. We estimated the overall evolutionary rate of the S. Gallinarum using BEAST. We applied the methodology described previously (60). The numbers of SNPs per year were determined by multiplying the evolutionary rates estimated with BEAST by the number of core SNP sites identified in the alignments. We hypothesize that a slower evolutionary rate in bacteria typically requires a lower SNP threshold when tracing transmission events using SNP distance analysis. Previous research found an average evolutionary rate of 1.97 SNPs per year (95% HPD, 0.48 to 4.61) across 22 different Salmonella serotypes. Our updated BEAST estimation for the evolutionary rate of S. Gallinarum suggests it is approximately 0.74 SNPs per year (95% HPD, 0.42 to 1.06). Based on these findings, as well as our previous experience with similar studies (61), we set a threshold of 2 SNPs (approximately representing less two years of evolution). Additionally, the geographic location was considered to enhance the precision of speculation regarding potential transmission sites and transmission direction.

MGE detecting

Earlier version: Four types of MGEs were detected: plasmids, transposons, integrons, and prophages. Plasmids were detected using Abricate v.1.0.1 with the PlasmidFinder (56) database. Only plasmids with a similarity of more than 90% and a coverage over 95% were identified. BacAnt (57) was used to detect integrons and transposons in the genomes. Only integrons or transposons with a similarity of more than 60% and a coverage greater than 60% will be identified. The prophages were detected using the Phaster pipeline. The genomic data were split into two temporary databases based on the number of contigs, one dataset containing single contig files and the other containing multiple contig files. The two databases were imported into the Phaster pipeline separately with default parameters.

Revised version: Four types of MGEs were detected: plasmids, transposons, integrons, and prophages. Plasmids were detected using Abricate v.1.0.1 with the PlasmidFinder (62) database. Only plasmids with a similarity of more than 90% and a coverage over 95% were identified. BacAnt (63) was used to detect integrons and transposons in the genomes. Only integrons or transposons with a similarity of more than 60% and a coverage greater than 60% will be identified. The prophages were detected using the Phaster pipeline. The genomic data were split into two temporary databases based on the number of contigs, one dataset containing single contig files and the other containing multiple contig files. The two databases were imported into the Phaster pipeline separately with default parameters.

Earlier version: Determination of the potential MGE carried ARG and horizontal ARG transfer frequency

Revised version: Determination of the MGE carried ARG and horizontal ARG transfer frequency

Earlier version: We have devised a pipeline to discern the coexistence of MGEs and ARGs. For plasmids and prophages, only ARGs situated within MGE regions are deemed MGE-carried. Regarding transposons and integrons, the search region was expanded to 5 kilobases upstream and downstream of ARGs, accounting for potential splicing errors. In this study, HGT events involving ARGs were defined as instances where ARGs exhibited perfect identity (100% coverage and 100% identity) and were located on MGEs across different strains. In a previous study (58), horizontal transfer of ARGs was identified based on gene regions with 100% coverage and over 99% identity. We believe that these stricter criteria in our pipeline more accurately reflect the transfer of ARGs within the S. Gallinarum population. HGT frequency was utilized to evaluate the extent of horizontal transfer of ARGs as follows: The pipeline developed in this study accepts Resfinder or Resistance Gene Identifier (RGI) results as input. Specific code and examples are uploaded to: https://github.com/tjiaa/Cal_HGT_Frequency/tree/main.

Revised version: We have devised a pipeline to discern the coexistence of MGEs and ARGs. For plasmids and prophages, only ARGs situated within MGE regions are deemed MGE-carried. Regarding transposons and integrons, the search region was expanded to 5 kilobases upstream and downstream of ARGs, accounting for potential splicing errors. In this study, HGT events involving ARGs were defined as instances where ARGs exhibited perfect identity (100% coverage and 100% identity) and were located on MGEs across different strains. In a previous study (64), horizontal transfer of ARGs was identified based on gene regions with 100% coverage and over 99% identity. We believe that these stricter criteria in our pipeline more accurately reflect the transfer of ARGs within the S. Gallinarum population. HGT frequency was utilized to evaluate the extent of horizontal transfer of ARGs as follows: The pipeline developed in this study accepts Resfinder or Resistance Gene Identifier (RGI) results as input. Specific code and examples are uploaded to: https://github.com/tjiaa/Cal_HGT_Frequency/tree/main.

Genomic data analysis

Earlier version: The serovars of all WGS genome data were confirmed through SISTR (59) v.1.1 and Seqsero2 (60). Multilocus Sequence Typing was carried out using MLST v.2.22 with the senterica_achtman_2 scheme. ARGs were detected by ResFinder (61). Both similarity and coverage were set to a minimum value of 90. The invasiveness index was calculated using methods previously reported (62). Specifically, Salmonella’s ability to cause intestinal or extraintestinal infections in hosts is related to the degree of genome degradation. We evaluated the potential for extraintestinal infection by 45 newly isolated S. Gallinarum strains (L2b and L3b) using a model that quantitatively assesses genome degradation. We analyzed each sample using the 196 top predictor genes for measuring the invasiveness of S. Gallinarum, employing a machine-learning approach that utilizes a random forest classifier and delta-bitscore functional variant-calling. This method evaluated the invasiveness of S. Gallinarum towards the host, and the distribution of invasiveness index values for each region was statistically tested using unpaired t-test. The code used for calculating the invasiveness index is available at https://github.com/Gardner-BinfLab/invasive_salmonella.

Revised version: The serovars of all WGS genome data were confirmed through SISTR (65) v.1.1 and Seqsero2 (66). Multilocus Sequence Typing was carried out using MLST v.2.22 with the senterica_achtman_2 scheme. ARGs were detected by ResFinder (67). Both similarity and coverage were set to a minimum value of 90. The invasiveness index was calculated using methods previously reported (32). Specifically, Salmonella’s ability to cause intestinal or extraintestinal infections in hosts is related to the degree of genome degradation. We evaluated the potential for extraintestinal infection by 45 newly isolated S. Gallinarum strains (L2b and L3b) using a model that quantitatively assesses genome degradation. We analyzed each sample using the 196 top predictor genes for measuring the invasiveness of S. Gallinarum, employing a machine-learning approach that utilizes a random forest classifier and delta-bitscore functional variant-calling. This method evaluated the invasiveness of S. Gallinarum towards the host, and the distribution of invasiveness index values for each region was statistically tested using unpaired t-test. The code used for calculating the invasiveness index is available at https://github.com/Gardner-BinfLab/invasive_salmonella.

Added in the revised version: ARGs vertical evolution control

Added in the revised version: In this study, HGTphyloDetect pipeline (68) was used to control for vertical evolution in the ARG sequences mentioned. We extracted base sequences for the eight ARGs shown in Figure 6b with an HGT frequency greater than zero (blaTEM-1B, sul1, dfrA17, aadA5, sul2, aph(3’’)-Ib, tet(A), aph(6)-Id). For blaTEM-1B, sul1, dfrA17, aadA5, and sul2, the HGT frequency reached 100% across different isolates, indicating that these ARG sequences have a unique ancestral sequence type. In contrast, due to the ResFinder settings requiring both similarity and coverage to meet a minimum value of 90%, the base sequences for aph(3’’)-Ib, tet(A), and aph(6)-Id are not unique. Consequently, we applied the HGTphyloDetect pipeline individually to each sequence type of ARGs to verify their association with HGT events. Specifically, among 436 bvSP isolates collected in China, we identified two sequence types of aph(3’’)-Ib, four sequence types of tet(A), and three sequence types of aph(6)-Id.

Added in the revised version: Subsequently, to identify potential ARGs horizontally acquired from evolutionarily distant organisms, we queried the translated amino acid sequences of each ARG against the National Center for Biotechnology Information (NCBI) non-redundant protein database. We then evaluated whether these sequences were products of HGT by calculating Alien Index (AI) scores and out_perc values. The calculation of Alien AI score is as follows: In this study, bbhG and bbhO represent the E-values of the best blast hit in ingroup and outgroup lineages, respectively. The outgroup lineage is defined as all species outside of the kingdom, while the ingroup lineage encompasses species within the kingdom but outside of the subphylum. Regarding the calculation method for out_perc: Finally, according to the definition provided by the HGTphyloDetect pipeline, ARGs with AI score ≥ 45 and out_perc ≥ 90% are presumed to be potential candidates for HGT from evolutionarily distant species.